# Introduction to Data Visualization


Welcome to the final and **the coolest** part of this course! 😎 I'm not saying that the previous topics aren't cool. But visualizations are the bridges that can connect your data analytics skills with the outside world. The last thing you want to do is print out a large DataFrame to report your findings.

Data visualization translates raw data into visual stories that reveal patterns, trends, and relationships. Humans process visual information far faster than text or tables—so effective visualizations help analysts think, communicate, and persuade.

> There is no such thing as information overload. There is only bad design.
>
> -- Edward Tufte


---

## From Data to Visualization: A Workflow

1. **Define the Question**: What decision or hypothesis are you exploring?
2. **Acquire and Clean Data**: Remove noise, handle missing values.
3. **Choose the Visual Form**: Match data type (categorical, continuous, time-series) to appropriate chart.
4. **Design and Annotate**: Add titles, captions, and labels.
5. **Iterate**: Test different layouts and get user feedback.


---

## Types of Data Visualization

### Exploratory vs. Explanatory Visualization

![Exploratory vs Explanatory](images/dataviz/exploratory-vs-explanatory.png)

| Type            | Goal                 | Audience     | Typical Tools                                     | Example                                           |
| --------------- | -------------------- | ------------ | ------------------------------------------------- | ------------------------------------------------- |
| **Exploratory** | Find insights        | Analyst      | Python (pandas, matplotlib, seaborn), R (ggplot2) | "What drives customer churn?"                     |
| **Explanatory** | Communicate insights | Stakeholders | Tableau, Power BI, Looker Studio                  | "Churn dropped 10% after loyalty program launch." |


### Data Patterns in Visualization

When we visualize data, we are not simply drawing charts - we are searching for _patterns_. Most quantitative insights can be grouped into five broad categories of data patterns. Recognizing these patterns helps analysts choose the most effective visual representation for their message.

![Data patterns](images/dataviz/data-patterns.png)

:::{seealso} Kevin Hartman's Digital Marketing Analytics: In Theory and In Practice

This categorization is from Kevin Hartman's [Digital Marketing Analytics: In Theory and In Practice](https://www.amazon.in/Digital-Marketing-Analytics-Theory-Practice/dp/B08J5BD5RM).

He taught digital marketing analytics in Gies iDegree programs, although he now teaches at his alma mater, Loyola University Chicago.

:::

### 1️⃣ Change

- **Definition:** Shows how a variable evolves over time.
- **Examples:** Line charts, area charts, or bar charts that track trends.
- **Use When:** You want to highlight growth, decline, or seasonal fluctuations.

### 2️⃣ Clustering

- **Definition:** Reveals natural groupings or segments within data.
- **Examples:** Scatter plots or bubble charts showing customer segments, product groups, or behavioral clusters.
- **Use When:** You want to explore differences and similarities between observations.

### 3️⃣ Relativity

- **Definition:** Displays how parts relate to a whole.
- **Examples:** Pie charts, donut charts, or stacked bar charts.
- **Use When:** You want to emphasize proportions or contribution to a total.

### 4️⃣ Ranking

- **Definition:** Compares ordered categories to identify leaders or laggards.
- **Examples:** Horizontal bar charts, lollipop charts, or sorted column charts.
- **Use When:** You want to show the top or bottom performers (e.g., top 10 sales regions).

### 5️⃣ Correlation

- **Definition:** Illustrates relationships between two or more quantitative variables.
- **Examples:** Scatter plots, correlation matrices, or regression lines.
- **Use When:** You want to determine whether changes in one variable are associated with changes in another.

### 💡 Why It Matters

Identifying these five data patterns helps analysts choose the right visualization for their story. Rather than focusing on chart types first, start by asking:

> "What kind of pattern am I trying to show - change, clustering, relativity, ranking, or correlation?"


### Common Chart Types

| Chart Type             | Best For                                      | Avoid When                |
| ---------------------- | --------------------------------------------- | ------------------------- |
| **Bar Chart**          | Comparing categorical values                  | Too many categories       |
| **Line Chart**         | Showing trends over time                      | Non-sequential categories |
| **Scatter Plot**       | Revealing relationships between two variables | Too few data points       |
| **Histogram**          | Showing data distribution                     | Comparing multiple groups |
| **Box Plot**           | Summarizing distribution & outliers           | Small samples             |
| **Pie / Donut Chart**  | Showing parts of a whole                      | Many small slices         |
| **Heatmap**            | Displaying matrix or correlation patterns     | Hard-to-read color scales |
| **TreeMap / Sunburst** | Hierarchical proportions                      | Need precise comparisons  |


---

## Common Pitfalls in Data Visualization

| Pitfall          | Description          | Better Practice           |
| ---------------- | -------------------- | ------------------------- |
| Overuse of 3D    | Distorts proportions | Stick to 2D               |
| Too Many Colors  | Confuses audience    | Use ≤ 5 meaningful colors |
| Truncated Y-Axis | Misleads differences | Start at 0 for bar charts |
| Dense Dashboards | Cognitive overload   | Prioritize key visuals    |
| Unlabeled Axes   | Ambiguous meaning    | Always label variables    |

:::{danger} Avoid 3D charts

There are very few cases where 3D charts are appropriate, such as a surface plot, or a spatial visualization. Otherwise, 3D charts often distort data and make it harder to interpret. Stick to 2D visualizations for clarity.

If you need to show multiple dimensions, consider using color, size, or facets instead of adding a third spatial dimension.

If you create a 3D chart because it looks "cool", I will track you down and personally make you redo it in 2D. 😡

:::


### 📊 Dataviz libraries

The mostly commonly used library for data visualization in introductory data analytics courses is [**matplotlib**](https://matplotlib.org/), a low-level visualization library for Python. [**seaborn**](https://seaborn.pydata.org/) is another popular library built on top of matplotlib that provides a higher-level interface for creating attractive and informative statistical graphics.

Below are some of the most popular and battle-tested data visualization libraries - all of which are free and open source:

- [**matplotlib**](https://matplotlib.org/): Low-level visualization library for Python
- [**seaborn**](https://seaborn.pydata.org/): High-level visualization library for Python **built on matplotlib**
- [**bokeh**](https://bokeh.org/): Interactive visualizations for modern web browsers
- [**plotnine**](https://plotnine.readthedocs.io/): Python implementation of [ggplot2](https://ggplot2.tidyverse.org/)
- [**plotly**](https://plotly.com/python/): Interactive visualization library supporting Python, JavaScript, and R
- [**altair**](https://altair-viz.github.io/): Declarative visualization library for Python based on [Vega](https://vega.github.io/vega/)

### Plot.ly

We'll use [plotly](https://plotly.com/python/), which provides both low-level and high-level interfaces to create publication-ready graphs.

![plotly logo image](images/dataviz/plotly-logo.png)


Import modules used by **🧭 Check Your Work** sections and the autograder.


```python
import unittest
import base64
import plotly

tc = unittest.TestCase()
```

---

## ✨ Visualize your dataset

▶️ First, run the code below to ensure you're using the correct version of plotly.


---

### 🎯 Pre-exercise: Import Packages


#### 👇 Tasks

- ✔️ Import the following Python packages.
  1. `pandas`: Use alias `pd`.
  2. `numpy`: Use alias `np`.
  3. `plotly.express`: Use alias `px`.
  4. `plotly.graph_objects`: Use alias `go`.


```python
# YOUR CODE BEGINS
import pandas as pd
import numpy as np
import plotly.graph_objects as go
import plotly.express as px

# YOUR CODE ENDS
```

#### 🧭 Check your work


<!-- grade_id: pre-exercise, points: 1, grade: true, solution: false -->
```python
import sys

tc.assertTrue(
    "pd" in globals(), "Check whether you have correctly imported Pandas with an alias."
)
tc.assertTrue(
    "np" in globals(), "Check whether you have correctly imported NumPy with an alias."
)

print(f"The current plotly version is {plotly.__version__}")
plotly_major_version = int(plotly.__version__.split(".")[0])
tc.assertGreaterEqual(
    plotly_major_version, 5, "Your plotly version should be greater than or equal to 5"
)

tc.assertIsNotNone(
    go.Figure,
    "Check whether you have correctly imported plotly.graph_objects with an alias go.",
)
tc.assertIsNotNone(
    px.scatter,
    "Check whether you have correctly imported plotly.express with an alias px.",
)
```

---

### 📌 Import dataset

Today, we work with an HR Dataset to uncover insights about HR metrics, measurement, and analytics. The data has been downloaded from [https://rpubs.com/rhuebner/hr_codebook_v14](https://rpubs.com/rhuebner/hr_codebook_v14) without any modification.


▶️ Run the code below to import an HR Dataset. 🐷👧👨🏻‍🦰👩🏼‍🦳👳🏽‍♂️👩🏾‍🦲🐼.


```python
# Display all columns
pd.set_option("display.max_columns", 50)

df_hr = pd.read_csv("https://github.com/bdi475/datasets/raw/main/HR-dataset-v14.csv")

display(df_hr)
```

---

## 📦 Box plot: summary of distribution

Box plots divide the data into 4 sections that each contain 25% of the data. It is useful to quickly identify the distribution of the data based on Q1, Q2 (median), and Q3.

![box plot explanation](https://github.com/bdi475/images/blob/main/lecture-notes/dataviz-python/box-plot-explanation-01.png?raw=true)


---

### 📌 GPA example

Before using the HR dataset, let's create a simple box plot of 12 different GPAs. NumPy is used here to calculate statistical figures.


```python
gpa = np.array(
    [3.33, 2.67, 3.0, 3.67, 3.67, 2.33, 3.0, 3.0, 2.67, 4.0, 3.33, 2.67, 4.0]
)
gpa
```

```python
fig = px.box(x=gpa, title="GPA Distribution (Horizontal Box Plot)")
fig.show()
```

```python
print(f"Mean: {np.mean(gpa)}")
print(f"Median: {np.median(gpa)}")
print(f"Q1: {np.quantile(gpa, 0.25)}")
print(f"Q3: {np.quantile(gpa, 0.75)}")
print(f"IQR: {np.quantile(gpa, 0.75) - np.quantile(gpa, 0.25)}")
```

### 🦄 Findings

- Median is `3`.
- Minimum is `2.33`.
- Maximum is `4`.
- Interquartile range is `1`.
  - You can calculate this value by subtracting Q1 from Q3: `3.67 - 2.67`.
- There is a positive skew.
  - This is also shown by comparing the mean and the median.


---

### 🎯 Exercise 1: Salary box plot (vertical)

#### 👇 Tasks

- ✔️ Draw a vertical box plot of `Salary` in `df_hr`.
- ✔️ Store your figure to a variable named `fig`.
- ✔️ Add an appropriate title to your figure.
- ✔️ Display the figure using `fig.show()`


```python
# YOUR CODE BEGINS
fig = px.box(df_hr, y="Salary", title="Salary Distribution (Vertical)")
fig.show()
# YOUR CODE ENDS
```

#### 🔑 Sample output

![Salary distribution vertical](https://github.com/bdi475/images/blob/main/exercises/intro-to-dataviz/salary_dispersion_box_plot_vertical.png?raw=true)


#### 🧭 Check your work


<!-- grade_id: exercise-01, points: 1, grade: true, solution: false -->
```python
tc.assertEqual(len(fig.data), 1, "There must be only one plot in your figure")
tc.assertIsNotNone(fig.layout.title.text, "Missing figure title")
tc.assertEqual(fig.data[0].type, "box", "Not a box plot")
tc.assertEqual(
    fig.data[0].orientation, "v", "Your plot should have a vertical orientation"
)
np.testing.assert_array_equal(fig.data[0].y, df_hr["Salary"], "Incorrect data")
```

---

### 🎯 Exercise 2: Salary box plot (horizontal)

#### 👇 Tasks

- ✔️ Draw a horizontal box plot of `Salary`.
- ✔️ Store your figure to a variable named `fig`.
- ✔️ Add an appropriate title to your figure.
- ✔️ Display the figure using `fig.show()`


```python
# YOUR CODE BEGINS
fig = px.box(df_hr, x="Salary", title="Salary Distribution (Horizontal)")
fig.show()
# YOUR CODE ENDS
```

#### 🔑 Sample output

![Salary distribution horizontal](https://github.com/bdi475/images/blob/main/exercises/intro-to-dataviz/salary_dispersion_box_plot_horizontal.png?raw=true)


#### 🧭 Check your work


<!-- grade_id: exercise-02, points: 1, grade: true, solution: false -->
```python
tc.assertEqual(len(fig.data), 1, "There must be only one plot in your figure")
tc.assertIsNotNone(fig.layout.title.text, "Missing figure title")
tc.assertEqual(fig.data[0].type, "box", "Not a box plot")
tc.assertEqual(
    fig.data[0].orientation, "h", "Your plot should have a horizontal orientation"
)
np.testing.assert_array_equal(fig.data[0].x, df_hr["Salary"], "Incorrect data")
```

---

### 🎯 Exercise 3: Salary distribution by citizenship status

#### 👇 Tasks

- ✔️ Draw horizontal box plots of `Salary` by `CitizenDesc`.
- ✔️ Store your figure to a variable named `fig`.
- ✔️ Add an appropriate title to your figure.
- ✔️ Display the figure using `fig.show()`


```python
# YOUR CODE BEGINS
fig = px.box(
    df_hr,
    x="Salary",
    y="CitizenDesc",
    title="Salary Distribution by Citizenship Status",
)
fig.show()
# YOUR CODE ENDS
```

#### 🔑 Sample output

![Salary distribution by citizenship status box plots](https://github.com/bdi475/images/blob/main/exercises/intro-to-dataviz/salary_dispersion_by_citizenship_status_box_plots_horizontal.png?raw=true)


#### 🧭 Check your work


<!-- grade_id: exercise-03, points: 1, grade: true, solution: false -->
```python
tc.assertEqual(len(fig.data), 1, "There must be only one plot in your figure")
tc.assertIsNotNone(fig.layout.title.text, "Missing figure title")
tc.assertEqual(fig.data[0].type, "box", "Not a box plot")
tc.assertEqual(
    fig.data[0].orientation, "h", "Your plot should have a horizontal orientation"
)
np.testing.assert_array_equal(fig.data[0].x, df_hr["Salary"], "Incorrect x-axis data")
np.testing.assert_array_equal(
    fig.data[0].y, df_hr["CitizenDesc"], "Incorrect y-axis data"
)
```

---

### 🎯 Exercise 4: Salary distribution by performance

#### 👇 Tasks

- ✔️ Draw horizontal box plots of `Salary` by `PerformanceScore`.
- ✔️ Store your figure to a variable named `fig`.
- ✔️ Add an appropriate title to your figure.
- ✔️ Display the figure using `fig.show()`


```python
# YOUR CODE BEGINS
fig = px.box(
    df_hr,
    x="Salary",
    y="PerformanceScore",
    title="Salary Distribution by Performance Score",
)
fig.show()
# YOUR CODE ENDS
```

#### 🔑 Sample output

![Salary distribution by performance score box plots](https://github.com/bdi475/images/blob/main/exercises/intro-to-dataviz/salary_dispersion_by_performance_score_box_plots_horizontal.png?raw=true)


#### 🧭 Check your work


<!-- grade_id: exercise-04, points: 1, grade: true, solution: false -->
```python
tc.assertEqual(len(fig.data), 1, "There must be only one plot in your figure")
tc.assertIsNotNone(fig.layout.title.text, "Missing figure title")
tc.assertEqual(fig.data[0].type, "box", "Not a box plot")
tc.assertEqual(
    fig.data[0].orientation, "h", "Your plot should have a horizontal orientation"
)
np.testing.assert_array_equal(fig.data[0].x, df_hr["Salary"], "Incorrect x-axis data")
np.testing.assert_array_equal(
    fig.data[0].y, df_hr["PerformanceScore"], "Incorrect y-axis data"
)
```

---

### 🎯 Exercise 5: Salary distribution by department

#### 👇 Tasks

- ✔️ Draw horizontal box plots of `Salary` by `Department`.
- ✔️ Store your figure to a variable named `fig`.
- ✔️ Add an appropriate title to your figure.
- ✔️ Set the height of your figure to `600`.
- ✔️ Display the figure using `fig.show()`


```python
# YOUR CODE BEGINS
fig = px.box(
    df_hr,
    x="Salary",
    y="Department",
    title="Salary Distribution by Department",
    height=600,
)
fig.show()
# YOUR CODE ENDS
```

#### 🔑 Sample output

![Salary distribution by department box plots](https://github.com/bdi475/images/blob/main/exercises/intro-to-dataviz/salary_dispersion_by_department_box_plots_horizontal.png?raw=true)


#### 🧭 Check your work


<!-- grade_id: exercise-05, points: 1, grade: true, solution: false -->
```python
tc.assertEqual(len(fig.data), 1, "There must be only one plot in your figure")
tc.assertIsNotNone(fig.layout.title.text, "Missing figure title")
tc.assertEqual(fig.data[0].type, "box", "Not a box plot")
tc.assertEqual(
    fig.data[0].orientation, "h", "Your plot should have a horizontal orientation"
)
tc.assertEqual(fig.layout.height, 600, "Incorrect height")
np.testing.assert_array_equal(fig.data[0].x, df_hr["Salary"], "Incorrect x-axis data")
np.testing.assert_array_equal(
    fig.data[0].y, df_hr["Department"], "Incorrect y-axis data"
)
```

---

## 🟪 Histogram: frequency distribution

Histograms display frequency distributions using bars of different heights.


```python
fig = px.histogram(x=np.random.randn(500))
fig.show()
```

---

### 🎯 Exercise 6: Salary histogram

#### 👇 Tasks

- ✔️ Draw a histogram of `Salary` in `df_hr`.
- ✔️ Store your figure to a variable named `fig`.
- ✔️ Add an appropriate title to your figure.
- ✔️ Display the figure using `fig.show()`


```python
# YOUR CODE BEGINS
fig = px.histogram(df_hr, x="Salary", title="Salary Distribution")
fig.show()
# YOUR CODE ENDS
```

#### 🔑 Sample output

![Salary distribution histogram](https://github.com/bdi475/images/blob/main/exercises/intro-to-dataviz/salary_histogram.png?raw=true)


#### 🧭 Check your work


<!-- grade_id: exercise-06, points: 1, grade: true, solution: false -->
```python
tc.assertEqual(len(fig.data), 1, "There must be only one plot in your figure")
tc.assertIsNotNone(fig.layout.title.text, "Missing figure title")
tc.assertEqual(fig.data[0].type, "histogram", "Not a histogram")
tc.assertEqual(
    fig.data[0].orientation, "v", "Your plot should have a vertical orientation"
)
np.testing.assert_array_equal(fig.data[0].x, df_hr["Salary"], "Incorrect data")
```

---

### 🎯 Exercise 7: Number of absences histogram

#### 👇 Tasks

- ✔️ Draw a histogram of `Absences` in `df_hr`.
- ✔️ Store your figure to a variable named `fig`.
- ✔️ Add an appropriate title to your figure.
- ✔️ Display the figure using `fig.show()`


```python
# YOUR CODE BEGINS
fig = px.histogram(df_hr, x="Absences", title="Number of Absence Distribution")
fig.show()
# YOUR CODE ENDS
```

#### 🔑 Sample output

![Number of absence histogram](https://github.com/bdi475/images/blob/main/exercises/intro-to-dataviz/number_of_absences_histogram.png?raw=true)


#### 🧭 Check your work


<!-- grade_id: exercise-07, points: 1, grade: true, solution: false -->
```python
tc.assertEqual(len(fig.data), 1, "There must be only one plot in your figure")
tc.assertIsNotNone(fig.layout.title.text, "Missing figure title")
tc.assertEqual(fig.data[0].type, "histogram", "Not a histogram")
tc.assertEqual(
    fig.data[0].orientation, "v", "Your plot should have a vertical orientation"
)
np.testing.assert_array_equal(fig.data[0].x, df_hr["Absences"], "Incorrect data")
```

---

### 🎯 Exercise 8: Salary histograms by gender

#### 👇 Tasks

- ✔️ Draw overlaid histograms of `Salary` in `df_hr` by `GenderID`.
- ✔️ Full code is provided below.

```python
# BEGIN SOLUTION
fig = go.Figure()
fig.add_trace(
    go.Histogram(
        x=df_hr[df_hr['GenderID'] == 0]['Salary'],
        name='Male'
    )
)

fig.add_trace(
    go.Histogram(
        x=df_hr[df_hr['GenderID'] == 1]['Salary'],
        name='Female'
    )
)

# Overlay both histograms
fig.update_layout(barmode='overlay')

# Reduce opacity to see both histograms
fig.update_traces(opacity=0.6)
fig.show()
# END SOLUTION
```


```python
# YOUR CODE BEGINS
fig = go.Figure()
fig.add_trace(go.Histogram(x=df_hr[df_hr["GenderID"] == 0]["Salary"], name="Male"))

fig.add_trace(go.Histogram(x=df_hr[df_hr["GenderID"] == 1]["Salary"], name="Female"))

# Overlay both histograms
fig.update_layout(barmode="overlay")

# Reduce opacity to see both histograms
fig.update_traces(opacity=0.6)
fig.show()
# YOUR CODE ENDS
```

#### 🧭 Check your work


<!-- grade_id: exercise-08, points: 1, grade: true, solution: false -->
```python
tc.assertEqual(len(fig.data), 2, "There must be two plots in your figure")
tc.assertEqual(fig.data[0].type, "histogram", "Not a histogram")
tc.assertEqual(fig.data[1].type, "histogram", "Not a histogram")
np.testing.assert_array_equal(
    fig.data[0].x, df_hr[df_hr["GenderID"] == 0]["Salary"], "Incorrect data"
)
np.testing.assert_array_equal(
    fig.data[1].x, df_hr[df_hr["GenderID"] == 1]["Salary"], "Incorrect data"
)
```